---
title: Chapter
description: 'The first chapter'
---

## Insert exercise title here

```yaml
type: VideoExercise
key: f197678c63
xp: 50
```

`@projector_key`
171805a72bdc5963d4532d41c5aec6bf

---

## A new title

```yaml
type: VideoExercise
key: ee85f12e42
xp: 50
```

`@projector_key`
f631216f10a9324607e9009741e598b2

---

## Insert exercise title here

```yaml
type: NormalExercise
key: e86d723676
xp: 100
```

this is the assignment

`@instructions`
- 1
- 2

`@hint`
a hint

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
