---
  title: "Chapter"
  description: "The first chapter"
  v2: true

---
## Machine Learning: What's the real challenge?

```yaml
type: VideoExercise 
xp: 50 
skills: 6
key: 58b94c3a47 
video_link: //player.vimeo.com/video/163565150 
video_hls: //videos.datacamp.com/transcoded/682_intro_to_ml/v1/hls-ch1_1.master.m3u8 
video_stream: https://player.vimeo.com/external/163565150.hd.mp4?s=7f21c9bf7b28967cfa8afb132999fa83788d26cc&profile_id=119
```


---
## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: 4ae51580cd   
```

demo

`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```

`@sample_code`

```{python}

```

`@solution`

```{python}

```

`@sct`

```{python}

```





